---
title       : Testing Spark
description : Testing Spark fo sho
attachments :
  slides_link : https://s3.amazonaws.com/assets.datacamp.com/course/teach/slides_example.pdf

--- type:NormalExercise lang:python xp:100 skills:1 key:2653a94a11
## A Spark Python exercise

This is an exercise to test spark for python. There's no pre exercise code.

*** =instructions

*** =pre_exercise_code
```{python}
```

*** =sample_code
```{python}
```

*** =solution
```{python}
```

*** =sct
```{python}
success_msg("Great work!")
```
