---
title       : Testing Spark for Python
description : Testing Spark for Python

--- type:NormalExercise lang:python xp:100 skills:1 key:2653a94a11
## A Spark Python exercise

On init, a spark cluster is started up.

*** =instructions
- You can experiment with spark in the IPython shell.

*** =pre_exercise_code
```{python}
```

*** =sample_code
```{python}
```

*** =solution
```{python}
```

*** =sct
```{python}
success_msg("Great work!")
```
